```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment1.ipynb")
```

# DS 453 / 653: Programming Assignment 1

**Due date**: Thursday, January 25 at 8pm on [Gradescope](https://www.gradescope.com/courses/710247).

_You must follow the Academic Code of Conduct and Collaboration Policy stated in the course syllabus at all times while working on this assignment._

This assignment contains 5 questions, each worth 1 point. You must receive at least 4 points to pass the assignment.

To begin, please execute the code block below:

```python
import otter
```

```python
grader = otter.Notebook()
```

**Question 1 (Integer to Hex):** Write a function `int_byte_to_hex` that receives a Python integer in the range 0 to 255, and returns a Python string corresponding to the same number in base 16 / hexadecimal format.

As background: remember that a _byte_ of information contains 8 bits, so it can be used to represent a number between 0 and 255 (inclusive). In crypto, we often convert between different representations of a byte of information. One common representation that we will use often in this class is hexadecimal format, in which each symbol (0-9 or a-f) corresponds to a number between 0 and 15, and a byte can be represented with exactly two symbols.

```python
def int_byte_to_hex(the_input):
    """ Read a single, decimal byte from the user and return a string of its
    hexidecimal value. This string should use lowercase and should always be
    exactly two characters long. Make sure you pad the beginning with a 0 if
    necessary, and make sure the string does NOT start with '0x'.

    Example test cases:

        255 -> "ff"
        10 -> "0a"
        65 -> "41"
        161 -> "a1"
    """
    assert(type(the_input) == int)               # the input is an integer
    assert(the_input >= 0 and the_input <= 255)  # between 0 and 255, inclusive

    # todo: implement this function
    hex_value = hex(the_input)[2:]
    return hex_value if len(hex_value) == 2 else "0" + hex_value

```

As you will see throughout this course, crypto code can be incredibly precise and finicky. For this reason: most programming questions will be accommpanied by **test cases**, which are example inputs and outputs to a function that allow you to check whether your code is working properly.

In this question, we provide test cases in human-readable format in the comment above, and we also provide unit tests for you to use below.

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2 (Binary string to Hex):** Write a function `string_to_hexstring` that receives a Python bytestring as input, and returns the hexadecimal representation of the same bytestring.

As a hint: the `binascii` library includes a function that will do this for you. Take a look at its documentation here: https://docs.python.org/3/library/binascii.html

```python
def string_to_hexstring(the_input):
    """ Take in a string (byte object), and return the hex string of the bytes
    corresponding to it. While the hexlify() command will do this for you,
    we ask that you instead solve this question by combining the methods
    you have written so far in this assignment.

    Example test case: b"puzzle" -> b"70757a7a6c65"
    """
    import binascii
    # todo: implement this function
    return binascii.hexlify(the_input)
```

Once again, we provide several unit tests for you. Note that in Python, bytestrings _must_ begin with the letter `b` before the quotation marks at the start of a string. Remember **always** to do this when creating bytestrings in this course! (Otherwise, you can find yourself stuck with an annoying error that is difficult to debug.)

The prefix `b` changes the type of the Python variable from `string` to `bytes`. We will usually prefer to use `bytes` in this course.

```python
print(type("puzzle"), type(b"puzzle"))
```

Output:
```
<class 'str'> <class 'bytes'>
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

**Question 3 (Length of bytestring vs hexadecimal string):** Note from the previous question that the length of a Python bytestring is different from the length of its hexadecimal encoding! Let's explore that in more detail in this question.

Write a function `compare_lengths` that receives an integer length, constructs a bytestring of the given length, and measures the length of the resulting hexadecimal string. (We've actually done the first part for you.)

```python
def hex_length(input_length):
    """ Construct a test_bytestring of the provided input_length,
        convert it to hexadecimal format, and then output
        the length of the resulting hex string.

    Example test case: 6 -> 12
    """
    import binascii
    test_bytestring = input_length * b'\x00'
    # todo: implement the rest of this function
    return len(binascii.hexlify(test_bytestring))
```

This time, we have only provided one publicly-visible test. But you should try to run this function on many inputs, and observe the resulting pattern.

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

**Question 4 (Hex to Base64):** Write a function `hexstring_to_base64` that receives a hexadecimal encoding of a bytestring, and returns the Base 64 encoding of the same bytestring.

(Note: this is actually the first challenge in the [Cryptopals](https://cryptopals.com/sets/1/challenges/1) crypto coding challenges. We might see some other challenges from Cryptopals later in the semester.)

This time, you might find the `base64` library to be useful. Whenever you explore a new Python library in this course, it is always a good idea to take a few minutes to read its documentation. Here's the documentation for `base64`: https://docs.python.org/3/library/base64.html

```python
import base64
import binascii

def hexstring_to_base64(the_input):
    """ Convert the_input from hex to base64 format.

    Example test cases:
    b'1de965a3ef96a2b95d' -> b'Hello++World'
    b'49276d206b696c6c696e6720796f757220627261696e206c696b65206120706f69736f6e6f7573206d757368726f6f6d' -> b'SSdtIGtpbGxpbmcgeW91ciBicmFpbiBsaWtlIGEgcG9pc29ub3VzIG11c2hyb29t'
    """
    # todo: implement the rest of this function
    return base64.b64encode(binascii.unhexlify((the_input)))
```

```python
print(hexstring_to_base64(b'1de965a3ef96a2b95d'))
```

Output:
```
b'Hello++World'
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

**Question 5 (One time pad):** Convert the inputs `hexMessage` and `hexKey` from hex to bytestrings, take the XOR of the two strings, and convert the result back to hex form.

In more detail: this procedure performs our first cryptographic function, a _one-time pad_. The one-time pad takes two inputs, typically called the `message` and the `key`, and returns an output called the `ciphertext`. We will explore the one-time pad in more detail in a future week. For now, all that matters is that the `ciphertext` is defined as `message` XOR `key` -- where the XOR is done on a byte-by-byte basis, and the XOR is done over the raw bytestrings.

However, it is often convenient for the inputs and outputs of a function to be hexadecimal representations of strings. The reason for that is simple: raw bytestrings might not always be printable, but hex strings always are. Consider for example the following 4-byte string.

```python
example_str = b'\x00\x01\x02\x03'
print("The length of example_str is:\t\t\t\t",         len(example_str))
print("If we try print it, the output is difficult to parse:\t",     example_str)
print("The hex representation is easy to read and copy:\t", string_to_hexstring(example_str))
```

Output:
```
The length of example_str is:				 4
If we try print it, the output is difficult to parse:	 b'\x00\x01\x02\x03'
The hex representation is easy to read and copy:	 b'00010203'
```

For this reason, we will encode our inputs in the hexadecimal representation. Your code must convert them to bytestrings first, and then perform the XOR function. There is a Python function in the `Crypto.Util` module that performs the string XOR for you.

(If you are running this code on your own machine rather than Google Colab: in order use this package, you must install the `pycrypto` package on your computer using a command such as `pip3 install pycrypto`.)

```python
from Crypto.Util.strxor import strxor
import base64
import binascii

def one_time_pad(hexMessage, hexKey):
    """ Convert both inputs from hexadecimal to raw bytestrings,
        take the xor of each byte of their strings, convert the
        result back into hexadecimal format, and output it.

        You can assume as a precondition that the two inputs have the same length.
        We provide some test cases below.
    """
    assert(len(hexMessage) == len(hexKey))
    # todo: implement the rest of this function
    message_str = binascii.unhexlify(hexMessage)
    key_str = binascii.unhexlify(hexKey)
    return binascii.hexlify(strxor(message_str, key_str))
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

## Submitting the Assignment

Congratulations on completing the first assignment! Here's how to submit it and receive credit.

**Documenting collaborators, sources, and AI tools:** In accordance with the collaboration policy, use the space below to report if you used any resources to complete this homework assignment, aside from the lecture notes and the course textbooks/videos. Specifically, please report:

1. Names of all classmates you worked with, and a short description of the work that you performed together.
2. All written materials that you used, such as books or websites (besides the lecture notes or textbooks). Please include links to any web-based resources, or citations to any physical works.
3. All code that you used from other sources. In particular, if you used an AI tool, then you must include the entire exchange with the AI tool, as per the [CDS Generative AI Assistance Policy](https://www.bu.edu/cds-faculty/culture-community/gaia-policy/).

_Your response:_

1. N/A

2. 

https://docs.python.org/3/library/base64.html#base64.b32encode,
https://docs.python.org/3/library/binascii.html, 
https://pycryptodome.readthedocs.io/en/latest/src/util/util.html#crypto-util-strxor-module

3. N/A

**Sending to Gradescope:**

After completing the assignment:
    if you did the assignment on Colab, download it in `.ipynb` format
    if you did the assignment locally on your machine, all you need to do is to find it in your directory

Then, submit only the `.ipynb` file to Coding Assignment 1 on the GradeScope and it takes a while for the auto grading system to check your work.

Tests for q1 (points: 1)
```
>>> int_byte_to_hex(255) == 'ff'
True
```
```
>>> int_byte_to_hex(10) == '0a'
True
```
```
>>> int_byte_to_hex(65) == '41'
True
```
```
>>> int_byte_to_hex(161) == 'a1'
True
```


Tests for q2 (points: 1)
```
>>> string_to_hexstring(b'puzzle') == b'70757a7a6c65'
True
```
```
>>> string_to_hexstring(b'hello-world') == b'68656c6c6f2d776f726c64'
True
```


Tests for q3 (points: 1)
```
>>> hex_length(6) == 12
True
```


Tests for q4 (points: 1)
```
>>> hexstring_to_base64(b'1de965a3ef96a2b95d') == b'Hello++World'
True
```
```
>>> hexstring_to_base64(b'49276d206b696c6c696e6720796f757220627261696e206c696b65206120706f69736f6e6f7573206d757368726f6f6d') == b'SSdtIGtpbGxpbmcgeW91ciBicmFpbiBsaWtlIGEgcG9pc29ub3VzIG11c2hyb29t'
True
```


Tests for q5 (points: 1)
```
>>> one_time_pad(b'001122', b'334455') == b'335577'
True
```
```
>>> one_time_pad(b'00112233', b'44556677') == b'44444444'
True
```
